```python
# Initialize Otter
import otter
grader = otter.Notebook("classification.ipynb")
```

```python
import numpy as np
import pandas as pd
import seaborn as sns
SEED = 2313
```

**Important note** Unless the instruction says otherwise, it's OK if your series or data frame results have additional rows/columns not requested, as long as the required information is shown.

The following dataset collects information about Google searches related to U.S. major sports leagues (5 professional leagues, plus college football and college baseball), tabulated by location in designated market areas (DMA). The values reflect the shares of all such searches attributed to each league. The last column is the share of the vote won by Donald Trump in the 2016 presidential election.

(Note how some extra parameters are used to designate the index column and header row within the CSV file.)


```python
fans = pd.read_csv("NFL_fandom_data-google_trends.csv",index_col=0,header=1)
fans
```

## Preliminaries

**1.** Overwrite `fans` with a frame in which the percent signs have been stripped out of all the columns, which are then converted to floating-point values.

```python
fans = ...
```

```python
grader.check("prelim-convert")
```

**2.** Create a feature matrix (in the form of a frame) from all the columns of `fans` except the last. 

```python
X = ...
```

```python
grader.check("prelim-features")
```

**3.** Make a Boolean label series that is `True` in each row where Trump got more than half of the vote.

```python
y = ...
```

```python
grader.check("prelim-labels")
```

**4.** Split the dataset, reserving 15% of it for testing and the rest for training. Make sure the split is random and starting from random state 3383.

```python
X_train,X_test,y_train,y_test = 
```

```python
grader.check("prelim-split")
```

## Decision tree

**1.** Train a decision tree of maximum depth 4 on the full dataset. Create a series of predictions using all the data, and compute the accuracy of the classifier.

**Important!** The decision tree classifier may randomly break ties among equivalent splits. In order to make your results reproducible, at the classifier creation you have to set its `random_state` equal to `SEED`, which is predefined above for you.

```python
dt = ...
yhat = ...
accuracy = ...
```

```python
grader.check("tree-accuracy")
```

**2.** Considering "Trump > 50%" to be the "positive" case, how many false positives does the classifier have?

```python
FP = ...
```

**3.** Considering "Trump > 50%" to be "positive," find the recall and F₁ score on the full dataset.

```python
recall = ...
F1 = ...
```

```python
grader.check("tree-scores")
```

**4.** Which of the sports leagues is most important to the classifier? What fraction of the overall impurity reduction does it account for?

```python
league = ...
fraction = ...
```

```python
grader.check("tree-importance")
```

## Nearest neighbors

**1.** Each row of the feature matrix `X` is a 7-dimensional vector. Find the 2-norm and 1-norm of the difference of the first two rows.

```python
two_norm = ...
one_norm = ...
```

```python
grader.check("knn-norm")
```

**3.** Train a nearest neighbors classifier with 5 neighbors on the training set, and measure its accuracy on the test set.

```python
accuracy = ...
```

```python
grader.check("knn-accuracy")
```

**4.** Which DMA most resembles a region that has 18% share in each of the professional leagues, plus 5% share in each of *CBB* and *CFB*?

```python
region = ...
```

```python
grader.check("knn-neighbors")
```

**5.** Repeat problem 3 above, but using standardized training data (i.e., Z-scores).

```python
accuracy = ...
```

## Support vector machine


**1.** Using the full dataset, train a linear kernel SVM classifier with $C=0.1$. Find the distance from the last sample observation to the separating hyperplane.

```python
def dot(u,v):
    return sum(u[i]*v[i] for i in range(len(u)))
dist = ...
```

```python
grader.check("svm-dist")
```


**2.** Using the train/test split, compute *F*₁ scores (with "Trump > 50%" as "positive") for a linear kernel SVM classifier with $C=1$, and for an RBF kernel classifier with $C=10$.

```python
linear = ...

rbf = ...
```

```python
grader.check("svm-f1")
```

**3.** Repeat the previous problem, but use standardized values (Z-scores).

```python
linear_standard = ...

rbf_standard = ...
```

```python
grader.check("svm-f1_standard")
```

## Validation and selection


**1.** Create a K-fold cross-validation scheme with 6 folds, starting at random state `SEED`. Use it to find a vector of balanced accuracy scores on the training data of a decision tree classifier with maximum depth 5.

**Important!** The decision tree classifier may randomly break ties among equivalent splits. In order to make your results reproducible, at the classifier creation you have to set its `random_state` equal to `SEED`, which is predefined above for you.

```python
test_score = ...
```

```python
grader.check("vaild-kfold")
```

<!-- BEGIN QUESTION -->

**2.** Using the same setup as in the previous problem, plot cross-validated training and test errors as a function of maximum decision tree depth running from 2 through 12.

```python

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**3.** Repeat the previous problem using k-nearest neighbor classifiers with the number of neighbors running from 2 through 20, applied to standardized scores.

```python

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**4.** Repeat the previous problem using an SVM classifier with RBF kernel and $C$ equal to $m/20$ as $m$ ranges over each integer from 1 through 20. Use standardized scores for the training.

```python

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**5.** For each of the three classifier types, select the hyperparameter value that gives the least error. For each classifier type, train on all of the training data at the optimum value, and then compute the confusion matrix on the test data.

```python

```

<!-- END QUESTION -->



---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Select *Kernel/Restart & Run All*, then save, then run this export cell again. Submit by pushing the resulting zip file to your GitHub assignment repo.

```python
grader.export(pdf=False, force_save=True)
```

Tests for knn-accuracy (points: None)
```
>>> assert 0.6 < accuracy < 0.7
```


Tests for knn-neighbors (points: None)
```
>>> assert X.loc[region,"NBA"] < 20 and X.loc[region,"CFB"] < 12
```
failure_message: Try again.


Tests for knn-norm (points: None)
```
>>> assert 12 < two_norm < 40
```
failure_message: 2-norm value is badly wrong.
```
>>> assert 30 < one_norm < 60
```
failure_message: Something is wrong with the numerical values.


Tests for knn-standard (points: None)


Tests for prelim-convert (points: None)
```
>>> assert all(fans.dtypes=="float64")
```
failure_message: Entries are not all numerical.
```
>>> assert list(fans.sum()) == [ 8093.,  4720.,  2814.,  1054.,  1112.,  985., 1921., 11287.55]
```
failure_message: Something is wrong with the numerical values.


Tests for prelim-features (points: None)
```
>>> assert type(X) == pd.DataFrame
```
failure_message: The feature matrix should have a DataFrame type.
```
>>> assert set(X.columns) == {'CBB', 'CFB', 'MLB', 'NASCAR', 'NBA', 'NFL', 'NHL'}
```
failure_message: The feature matrix has the wrong columns.
```
>>> assert X.shape == (207,7)
```
failure_message: The shape is incorrect.


Tests for prelim-labels (points: None)
```
>>> assert y.dtype=='bool'
```
failure_message: The result should be a series with Boolean values.
```
>>> assert sum(y)==138
```
failure_message: There should be 138 True values.


Tests for prelim-split (points: None)
```
>>> assert X_train.shape[0]==175 and X_test.shape[0]==32
```
failure_message: The subset sizes are incorrect.
```
>>> assert not all(X.NBA[:10].values == X_train.NBA[:10].values)
```
failure_message: The ordering was not randomized.
```
>>> assert all(np.isclose(X_train.iloc[0,:].values,[41,24,18,10,3,3,3]))
```
failure_message: Check the random state argument.


Tests for svm-dist (points: None)
```
>>> assert 4 < dist < 10
```


Tests for svm-f1 (points: None)
```
>>> assert linear > rbf and linear > 0.85 and rbf > 0.8
```


Tests for svm-f1_standard (points: None)
```
>>> assert rbf_standard > linear_standard and rbf_standard > 0.9 and linear_standard > 0.85
```


Tests for tree-FP (points: None)


Tests for tree-accuracy (points: None)
```
>>> assert type(dt)==DecisionTreeClassifier and dt.max_depth==4
```
failure_message: You must use a decision tree of maximum depth 4.
```
>>> assert 0.80 < accuracy < 0.85
```
failure_message: The accuracy is not in the correct range.
```
>>> assert sum(yhat)==140
```
failure_message: There is an incorrect number of True predictions.


Tests for tree-importance (points: None)
```
>>> assert league in X.columns
```
failure_message: The league must be the name of a data column.
```
>>> assert 0 <= fraction <= 1
```
failure_message: The fraction must be a float between 0 and 1.


Tests for tree-scores (points: None)
```
>>> assert 0.82 < recall < 0.90
```
failure_message: The recall is not in the correct range.
```
>>> assert 0.82 < F1 < 0.90
```
failure_message: The F1 value is not in the correct range.


Tests for vaild-kfold (points: None)
```
>>> assert all(0.55 < test_score) and all(test_score < 0.84)
```
```
>>> assert np.isclose(test_score.mean(),0.6934144072301965)
```
