```python
# Initialize Otter
import otter
grader = otter.Notebook("regression.ipynb")
```

```python
import numpy as np
import pandas as pd
import seaborn as sns
SEED = 3383
```

## 1. Regression

This dataset contains information on ticket prices for flights between six cities in India:

```python
flights = pd.read_csv("flights.csv",index_col=0)
flights.head(5)
```

1.1. Create a boolean vector or series called `trips` that is `True` only for flights with class *Economy*, source city *Delhi*, and destination city *Mumbai*.

```python
trips = ...
```

```python
grader.check("flights-trips")
```

1.2. Prepare a feature matrix `X` for the rows in `trips` and the columns *departure_time*, *airline*, *duration*, *days_left*, and *stops*. In the *stops* column, replace `"zero"` with 0, `"one"` with 1, and `"two_or_more"` with 2.
   
Prepare a feature vector `y` for *price* on the same rows.

```python
X = ...
y = ...
```

```python
grader.check("flights-features")
```

1.3. Perform a linear regression for the price with the predictors *days_left* and *duration*. Find the coefficient of determination for the fit.

```python
CofD = ...
```

```python
grader.check("regress-twovars")
```

1.4. Does increasing *days_left* cause the price to increase, or decrease? Does increasing *duration* cause the price to increase, or decrease? Answer `True` or `False` for each variable.

```python
# True or False in each case
increasing_days_left_increases_price = ...
increasing_duration_increases_price = ...
```

1.5. Create a new frame `Xdum` that replaces the *airline* and *departure_time* features with dummy variables. **Use `drop_first=True` for when creating the dummies.** This option replaces a category of $k$ unique values with $k-1$ variables, leaving out redundancy.

Then retrain the linear regressor using `Xdum` and compute the new coefficient of determination.

```python
Xdum = ...
CofD_dum = ...
```

```python
grader.check("regress-dummies")
```

1.6. Which airline, when chosen with all else being unchanged, tends to cause the greatest increase in the price? All else being equal, what is the best departure time for lowering the price?

```python
# Use a string value as the answer:
airline_biggest_increase = ...
time_biggest_decrease = ...
```

```python
grader.check("regress-effects")
```

1.7. Use LASSO on the feature matrix with dummies and increase the regularization parameter until one of the coefficients is dropped. Which column corresponds to the dropped coefficient?

```python
first_dropped = ...
```

```python
grader.check("regress-lasso")
```

1.8. Use a decision tree regressor on the feature matrix with dummy variables, and compute its coefficient of determination score. Also, determine which feature is deemed to be most important by the regressor.

```python
CofD_dtr = ...
top_feature = ...
```

```python
grader.check("flight-dtree")
```

## 3. Probabilistic classification

This dataset contains the results of passenger satisfaction surveys on U.S. airlines.

```python
satisfaction = pd.read_csv("passenger_satisfaction.csv",index_col=0).dropna()
satisfaction.head(6)
```

Some of the columns are features that can be deduced independently of surveys, at least in principle. We will separate the two types of features for the analysis that follows.

```python
objective = ['Gender', 'Customer Type', 'Age', 'Class',
        'Flight Distance','Departure Delay in Minutes', 'Arrival Delay in Minutes',]
subjective = ['Type of Travel','Inflight wifi service',
       'Departure/Arrival time convenient', 'Ease of Online booking',
       'Gate location', 'Food and drink', 'Online boarding', 'Seat comfort',
       'Inflight entertainment', 'On-board service', 'Leg room service',
       'Baggage handling', 'Checkin service', 'Inflight service',
       'Cleanliness']
target = 'satisfaction'
```

3.1 Create a feature matrix `XO` for just the `objective` features, with dummy variables, again using `drop_first=True`. Also make a label series `y` for the *satisfaction* column. Split off 20% of the data into a test set, using a shuffle with random state `SEED`.

```python
XO = ...
ys = ...
...
```

```python
grader.check("satisfy-objective")
```

3.2 Using the objective-only training data, perform a logistic regression with the option `penalty="none"`, which disables regularization. Make sure to use a pipeline with column standardization. 

Using the test set, find the AUC-ROC score of the regression.

```python
lr_obj = ...
AUC_obj = ...
```

```python
grader.check("satisfy-objective-logr")
```

<!-- BEGIN QUESTION -->

3.3. Plot the ROC curve of the regressor for detection of the state *satisfied* for **the first 1000 members** of the test set.

```python

```

<!-- END QUESTION -->

3.4 Now use the objective and subjective features together, converting to dummies and splitting into train/test like before. 

Perform logistic regression with $C=0.01$, `penalty="l1"` (i.e., LASSO), and `solver="liblinear"` (since the default will not work with LASSO). Compute the AUC-ROC score.

```python
XA = ...
...

lr_all = ...
AUC_all = ...
```

```python
grader.check("satisfy-all-auc")
```

3.5. Use the coefficients of the regression (with all features) to answer the following questions. 

(a) Which feature or features are dropped by the regularization? (Answer with a **list** of strings.)

(b) Which of the subjective features has the greatest positive effect on satisfaction? (Answer with a string.)

(c) Which type of travel tends to produce greater satisfaction? (Answer with string `"Personal"` or `"Business"`.)

```python
dropped_features = ...
greatest_subjective = ...
satisfying_travel = ...
```

```python
grader.check("satisfy-all-factors")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Select *Kernel/Restart & Run All*, then save, then run this export cell again. Submit by pushing the resulting zip file to your GitHub assignment repo.

```python
grader.export(pdf=False, force_save=True)
```

Tests for flight-dtree (points: None)
```
>>> assert 0.7 < CofD_dtr < 0.75
```
failure_message: Regressor score should be in the interval (0.7,0.75).
```
>>> assert top_feature in Xdum.columns
```
failure_message: Top feature should be a column name in Xdum.


Tests for flights-features (points: None)
```
>>> assert set(X.columns) == {'airline', 'days_left', 'departure_time', 'duration', 'stops'}
```
failure_message: The feature names are incorrect.
```
>>> assert np.isclose(X["days_left"].mean(),25.915948707673813)
```
failure_message: You do not have the correct feature rows.
```
>>> assert np.isclose(X["stops"].mean(),0.7476457623722701)
```
failure_message: The "stops" column is not numeric, or has the wrong values.
```
>>> assert np.isclose(y.mean(),6059.826086956522)
```
failure_message: You do not have the correct label rows.


Tests for flights-trips (points: None)
```
>>> assert sum(trips)==9982
```
failure_message: There should be 9982 trips selected.


Tests for regress-dummies (points: None)
```
>>> assert CofD_dum > 0.5 and CofD_dum < 0.55
```
failure_message: The result should be in the range (0.5,0.55).


Tests for regress-effects (points: None)
```
>>> assert airline_biggest_increase in X["airline"].unique()
```
failure_message: Your answer must be one of the airlines.
```
>>> assert time_biggest_decrease in X["departure_time"].unique()
```
failure_message: Your answer must be one of the airlines.


Tests for regress-increase (points: None)


Tests for regress-lasso (points: None)
```
>>> assert first_dropped in Xdum.columns
```
failure_message: Wrong!!


Tests for regress-twovars (points: None)
```
>>> assert CofD > 0.36 and CofD < 0.40
```
failure_message: Result should be in the range (0.36,0.40).


Tests for satisfy-all-auc (points: None)
```
>>> assert 0.9 < AUC_all < 0.93
```
failure_message: The AUC score should be in the interval (0.9,0.93).
```
>>> assert len(XA.columns) == 23
```
failure_message: The columns of the XA feature matrix are incorrect.


Tests for satisfy-all-factors (points: None)
```
>>> assert type(dropped_features)==list and set(dropped_features).difference(set(XA.columns)) == set()
```
failure_message: Answer (a) with a list of strings from the full column name set.
```
>>> assert type(greatest_subjective)==str and greatest_subjective in subjective
```
failure_message: Answer (b) with a string from the objective column names.
```
>>> assert satisfying_travel in ["Personal","Business"]
```
failure_message: Answer (ac with either "Personal" or "Business".


Tests for satisfy-objective (points: None)
```
>>> assert set(XO.columns) == set(['Age', 'Flight Distance', 'Departure Delay in Minutes',
...        'Arrival Delay in Minutes', 'Gender_Male',
...        'Customer Type_disloyal Customer', 'Class_Eco', 'Class_Eco Plus'])
```
failure_message: Your feature matrix does not have the right columns.
```
>>> assert len(ys)==129487 and sum(ys=="satisfied")==56262
```
failure_message: Something is wrong with the label vector.
```
>>> assert np.isclose(XO_train["Age"].mean(),39.43230458832502)
```
failure_message: Something is wrong with the training matrix.


Tests for satisfy-objective-logr (points: None)
```
>>> assert 0.78 < AUC_obj < 0.80
```
failure_message: AUC score should be in the interval (0.78,0.80).
