```python
# Initialize Otter
import otter
grader = otter.Notebook("basics.ipynb")
```

If you want to run the solution checker, and you have never installed the grader package, uncomment and run the cell below. (It won't hurt to run if you have installed before.)

```python
#! pip install otter-grader
```

```python
import pandas as pd
import numpy as np
import seaborn as sns
```

Take note that **passing the grader checks does not mean that your answer is correct!** The grader performs some tests to catch significant problems, but does not validate against the true solutions.

## Enter the matrix

**1.** Create a vector of the first 8 prime numbers.

```python
primes = ...
```

```python
grader.check("vector-primes")
```

**2.** Create the vector [1, 4, 9, 16, …, 1521, 1600]. 

```python
squares = ...
```

```python
grader.check("vector-squares")
```

**3.** Create a matrix with 4 rows and 2 columns, with all entries equal to 2.5.

```python
C = ...
```

```python
grader.check("matrix-constant")
```

**4.** Write a function that accepts a matrix as input and returns the matrix that results from subtracting within every column the maximum value in that column.

```python
def subtract_max(A):
    ...
```

```python
grader.check("matrix-subtract")
```

## School daze

**1.** Create a series called `schools`, that has series name `"School"` and the values `"Delaware"`, `"Penn"`, `"Drexel"`, and `"Villanova"`.

```python
schools = ...
```

```python
grader.check("school-series")
```

**2.** Make a data frame called `uni` that looks like this:

![data frame](df_snapshot.png)

Note: The years should be of `int64` data type.

```python
uni = ...
```

```python
grader.check("school-df")
```

## Fight, fight, fight!

<!-- BEGIN QUESTION -->

**1.** Load the `fight-songs.csv` file as a data frame called `songs`. Display the first 5 rows.

```python
songs = ...
...
```

<!-- END QUESTION -->

**2.** Select the *song_name* column and assign it to a variable called `names`. Define a series called `subset` having the 10th through 18th values of `names`.

```python
names = ...
subset = ...
```

```python
grader.check("fight-names")
```

**3.** What is the maximum value in the *trope_count* column?

```python
max_tropes = ...
```

**4.** Select from `songs` only the rows having an index value divisible by 10, and only the columns *year* and *official_song*. 

```python
selection = ...
```

```python
grader.check("fight-select")
```

**5.** Create a series showing the number of rows in `songs` belonging to each conference.

```python
num_conf = ...
```

```python
grader.check("fight-numconf")
```

**6.** Make a data frame selecting all the songs written in 1915.

```python
songs1915 = ...
```

```python
grader.check("fight-1915")
```

**7.** How many of the fight songs last less than one minute?

```python
num_short = ...
```

<!-- BEGIN QUESTION -->

**8.** Plot the distribution of *sec_duration* as a histogram over 12 bins.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**9.** Plot histograms of *bpm*, separated into columns by the value of *student_writer*.

```python
...
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**10.** Make a violin plot showing the distribution of *bpm* for each *conference*.

```python
...
```

<!-- END QUESTION -->

**11.** Create a new data frame that replaces in `songs` all instances of `"Yes"` with `True` and all instances of `"No"` with `False`.

```python
songsTF = ...
```

```python
grader.check("fight-replace")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Select *Kernel/Restart & Run All*, then save, then run this export cell again. Submit by pushing the resulting zip file to your GitHub assignment repo.

```python
grader.export(pdf=False, force_save=True)
```

Tests for fight-1915 (points: None)
```
>>> songs1915.shape == (4,23)
True
```
failure_message: The value does not have the correct shape.
```
>>> set(songs1915.year) == {"1915"}
True
```
failure_message: The year column is not always correct.


Tests for fight-60sec (points: None)


Tests for fight-max (points: None)


Tests for fight-names (points: None)
```
>>> type(names) == pd.Series
True
```
failure_message: names is not of type Series.
```
>>> len(names) == 65
True
```
failure_message: names has the wrong length.
```
>>> type(subset) == pd.Series
True
```
failure_message: subset is not of type Series.
```
>>> len(subset) == 9
True
```
failure_message: subset has the wrong length.
```
>>> subset.values
array(['Fight, Raiders, Fight', 'Hail West Virginia', 'Oskee-Wow-Wow',
       'Indiana, Our Indiana', 'Iowa Fight Song', 'Maryland Fight Song',
       'The Victors', 'Victory for MSU', 'The Minnesota Rouser'],
      dtype=object)
```


Tests for fight-numconf (points: None)
```
>>> type(num_conf) == pd.Series
True
```
failure_message: The value is not of type Series.
```
>>> len(num_conf) == 6
True
```
failure_message: The value has the wrong length.
```
>>> num_conf.name == "conference"
True
```
failure_message: The Series name property is incorrect.


Tests for fight-replace (points: None)
```
>>> songs.shape == (65,23)
True
```
failure_message: The value does not have the correct shape.
```
>>> songsTF.nonsense.dtype == "bool"
True
```
failure_message: Some values were not converted to True/False.


Tests for fight-select (points: None)
```
>>> np.all(selection.index%10 == 0)
True
```
failure_message: Not all row indices are divisible by 10.
```
>>> len(selection) == 1 + len(songs)//10
True
```
failure_message: Result has the wrong length.
```
>>> list(selection.columns) == ["year","official_song"]
True
```
failure_message: The frame's columns are incorrect.


Tests for matrix-constant (points: None)
```
>>> type(C) == np.ndarray
True
```
failure_message: The value is not of array type.
```
>>> assert C.shape == (4,2)
```
failure_message: The value does not have the correct shape.


Tests for matrix-subtract (points: None)
```
>>> testmat = np.array([[1,3,4],[5,0,2]])
>>> np.all( np.isclose(subtract_max(testmat),np.array([[-4,0,0],[0,-3,-2]]) ) )
True
```
failure_message: Incorrect result.


Tests for school-df (points: None)
```
>>> type(uni) == pd.DataFrame
True
```
failure_message: The value is not a DataFrame.
```
>>> list(uni.columns) == ['Founded', 'Mascot']
True
```
failure_message: The frame does not have the correct column names.


Tests for school-series (points: None)
```
>>> type(schools) == pd.Series
True
```
failure_message: The value is not of type Series.
```
>>> schools.name == "School"
True
```
failure_message: The Series name property is not correct.
```
>>> np.all(schools.values == ["Delaware","Penn","Drexel","Villanova"])
True
```
failure_message: The Series does not have correct values.


Tests for vector-primes (points: None)
```
>>> type(primes) == np.ndarray
True
```
failure_message: The value is not of array type.
```
>>> assert primes.shape == (8,)
```
failure_message: The value does not have the correct shape.


Tests for vector-squares (points: None)
```
>>> type(squares) == np.ndarray
True
```
failure_message: The value is not of array type.
```
>>> assert squares.shape == (40,)
```
failure_message: The value does not have the correct shape.
