```python
# Initialize Otter
import otter
grader = otter.Notebook("network lab.ipynb")
```

```python
import numpy as np
import pandas as pd
import seaborn as sns
import networkx as nx
SEED = 3383
```

# 0. Data preparation

The following is a network of connections within part of the brain of a fly.

```python
fly = nx.read_edgelist("bn-fly-drosophila_medulla_1.edges",nodetype=int)
n = fly.number_of_nodes()
print(n,"nodes")
```

This network has a feature that we have not allowed so far: some nodes link to themselves. These appear, for example, as nonzeros on the diagonal of the adjacency matrix. 

```python
A = nx.adjacency_matrix(fly).toarray()
print(sum(A[i,i] for i in range(n)),"self-linked nodes")
```

We will get rid of these self-links by zeroing out the diagonal and reconstructing the graph from the modified adjacency matrix. This will also relabel the nodes to go in order from 0 to $n-1$.

```python
np.fill_diagonal(A,0)
fly = nx.from_numpy_array(A)
nodes = pd.Index(fly.nodes)
```

# 1. Basic info

## 1.1 

Compute the average node degree of the network.

```python
kbar = ...

# No need to change this line.
print("average degree is",kbar)
```

```python
grader.check("info-kbar")
```

<!-- BEGIN QUESTION -->

## 1.2 

Draw the ego graph of node 80, with labels for the nodes.

```python

```

<!-- END QUESTION -->

## 1.3 

Create a list of all the nodes that are adjacent to node 9.

```python
nbrs9 = ...

# No need to change the following line.
print("neighbors of node 9:",nbrs9)
```

```python
grader.check("info-neighbors")
```

# 2. Clustering and distance

## 2.1
Create a Series for the local clustering coefficients in the network.

```python
cluster = ...

# No need to change the following lines.
print("clustering coeffs:")
print(cluster.describe())
```

```python
grader.check("cluster-values")
```

## 2.2

Compute $T(110)$, the number of triangles that have node 110 as a vertex. Then, apply the formula

$$
\frac{2T(110)}{d_{110}(d_{110}-1)},
$$

where $d_{110}$ is the degree of node 110, to get the clustering coefficient of 110.

```python
# Set this to the number of triangles.
T110 = ...

# Set this to the local clustering coefficient.
C110 = ...

# No need to change the following line.
print(T110,"adjacent triangles, giving clustering coefficient",C110)
```

```python
grader.check("cluster-110")
```

## 2.3

Let $k$ be the integer closest to `kbar`. For each $q=0,0.02,0.04,0.06,\ldots$, compute and print out the average clustering coefficient of a Watts–Strogatz graph for $n$, $k$, and $q$, with random seed equal to `SEED`. When the average coefficient is smaller than the average clustering coefficient of the `fly` graph, stop.

```python
# Exit this loop when the average WS clustering exceeds the average clustering in the fly graph.
for q in np.arange(0,0.5,0.02):

# No need to change the following line.
print("final value of q is",q)
```

```python
grader.check("cluster-WS")
```

## 2.4
Compute the mean shortest path length between all pairs of nodes in the largest connected component of the `fly` graph. (This computation might take 20-60 seconds, depending on the speed of the computer you use.)

```python
avg_dist = ...

# No need to change the following line.
print("average shortest path length:",avg_dist)
```

```python
grader.check("distance-average")
```

# 3. Degree distribution 

## 3.1
Create a Series of the node degrees in the `fly` graph.

```python
degrees = ...

# No need to change the following line.
print("node degrees for fly graph:\n",degrees.describe())
```

```python
grader.check("degree-series")
```

<!-- BEGIN QUESTION -->

## 3.2
 Using log scales on both axes, plot a histogram of the degree distribution. You should see a rough linear trend in the resulting plot.

```python

```

<!-- END QUESTION -->

## 3.4
 Create a Barabási–Albert graph with random seed `SEED`, the same number of nodes as the `fly` graph, and parameter $m$ equal to the integer nearest to half the average degree of the `fly` graph. Find the standard deviation of its degree distribution.

```python
# This should be assigned to a Graph object.
BA = ...

# This should get the value of the standard deviation of the degree distribution.
BA_deg_stdev = ...

# No need to change the following lines.
print("BA graph has",len(BA),"nodes and",BA.number_of_edges(),"edges")
print("and node degrees with standard deviation",BA_deg_stdev)
```

```python
grader.check("degree-BA")
```

# 4. Centrality

## 4.1
Create a data frame indexed by node and with columns named *degree*, *betweenness*, and *eigenvector* for centrality measures.

```python
# This should be a data frame indexed by nodes and with 3 columns for centrality measures. Fidn the mean value in each column.
centrality = ...
# These should be assigned to the means of the three columns.
dc_mean = ...
bc_mean = ...
ec_mean = ...

# No need to change the following lines.
print("centrality scores\n:",centrality)
print("with means",dc_mean,",",bc_mean,",",ec_mean)
```

```python
grader.check("central-columns")
```

## 4.2
Of the three pairings degree-betweenness, degree-eigenvector, and betweenness-eigenvector, find the Pearson correlation coefficient of centrality scores for the pairing that is the *least* strongly correlated.

```python
min_correlation = ...

# No need to change the following line.
print("weakest correlation is",min_correlation)
```

```python
grader.check("central-corr")
```

## 4.3
For each type of centrality score, find the 10 highest-scoring nodes.

```python
# Each should be an array, list, or Series.
top_degree = ...
top_betweenness = ...
top_eigenvector = ...

# No need to change these lines.
results = pd.DataFrame({"degree":top_degree,"betweenness":top_betweenness,"eigenvector":top_eigenvector})
print(results)
```

```python
grader.check("central-rank")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Select *Kernel/Restart & Run All*, then save, then run this export cell again. Submit by pushing the resulting zip file to your GitHub assignment repo.

```python
grader.export(pdf=False, force_save=True)
```

Tests for central-columns (points: 4)
```
>>> assert all(centrality.index == fly.nodes)
```
failure_message: The index of the frame has to be the nodes property of the fly graph.
```
>>> assert centrality.shape == (fly.number_of_nodes(),3)
```
failure_message: The frame has the wrong shape.
```
>>> assert np.isclose(centrality.betweenness.mean(),0.001060709032607)
>>> assert np.isclose(centrality.eigenvector.mean(),0.0131148,rtol=1e-4)
```
failure_message: Some of the frame's values are wrong.


Tests for central-corr (points: 4)
```
>>> assert 0 <= min_correlation <= 1
```
failure_message: The correlation coefficient should be between 0 and 1.


Tests for central-rank (points: None)
```
>>> assert sum((results==133).sum()) == 2
```
failure_message: Node 133 should appear in two of the lists.


Tests for cluster-110 (points: 4)
```
>>> assert 580 <= T110 <= 600
```
failure_message: T110 should be in the interval [580,600].
```
>>> assert 0.14 <= C110 <= 0.16
```
failure_message: C110 should be in the interval [0.14,0.16].


Tests for cluster-WS (points: 4)
```
>>> assert 0.2 <= q <= 0.3
```
failure_message: The final q should be in the interval $[0.2,0.3]$.


Tests for cluster-values (points: 4)
```
>>> assert np.isclose(cluster.mean(),0.26284869232685665)
```
failure_message: The clustering values are wrong.


Tests for degree-BA (points: 4)
```
>>> assert np.isclose(nx.average_clustering(BA),0.02941239985)
```
failure_message: The BA graph is wrong.


Tests for degree-series (points: 2)
```
>>> assert degrees.count()==1781
```
failure_message: The largest connected component has 1770 nodes.
```
>>> assert 9.5 <= degrees.mean() <= 10.5
```
failure_message: The average degree is in the interval $[9.5,10.5]$.


Tests for distance-average (points: 4)
```
>>> assert 2.8 <= avg_dist <= 3.2
```
failure_message: The average distance is in the interval [2.8,3.2].


Tests for info-kbar (points: 2)
```
>>> assert 9 <= kbar <= 12
```
failure_message: The average degree should be in the interval [9,12].


Tests for info-neighbors (points: 4)
```
>>> assert len(nbrs9) == fly.degree[9]
```
failure_message: The number of neighbors should equal the degree of the node.
