```python
# Initialize Otter
import otter
grader = otter.Notebook("stats.ipynb")
```

```python
import numpy as np
import pandas as pd
import seaborn as sns
```

**Important note** Unless the instruction says otherwise, it's OK if your series or data frame results have additional rows/columns not requested, as long as the required information is shown.

## Waterfowl surveys


**1.** Load the dataset `Aerial_Waterfowl_Survey_Data.csv` as a data frame called `birds`. (This dataset shows counts of various species of waterfowl observed during aerial surveys.)

```python
birds = ...
```

```python
grader.check("birds-load")
```

**2.** Create a series for the total number of birds observed in Zone 3 for each species. 

```python
zone3 = ...
```

```python
grader.check("birds-zone3")
```

**3.** Make a series of the total number of Canada geese observed in each zone.

```python
canada = ...
```

```python
grader.check("birds-canada")
```

<!-- BEGIN QUESTION -->

**4.** Make a box plot, grouped by month, of the Mallard counts.

```python
...

```

<!-- END QUESTION -->

**5.** Count the total number of outliers for December in the previous plot.

```python
outliersDec = ...
```

<!-- BEGIN QUESTION -->

**6.** Make a data frame for just the species American black duck, Mallard, Wood duck, and Ring-necked duck. Create a grid of pairwise scatter plots for the duck counts.

```python
ducks = ...
...
```

<!-- END QUESTION -->

**7.** Make a frame of the Pearson correlation coefficients between all pairings of the four duck species.

```python
duckduck = ...
```

```python
grader.check("birds-pearson")
```

**8.** Repeat the previous question, but using Spearman coefficients. 

```python
duckspear = ...
```

```python
grader.check("birds-spearman")
```

## Restaurant inspections



**1.** Load the dataset `Restaurant_Inspection_Violations.csv` as a data frame. (This dataset has the results of recent restaurant inspections in Delaware.)

```python
inspects = ...
```

```python
grader.check("inspect-load")
```

**2.** A `NaN` value in the `Violation Code` column seemingly indicates that no violation was found. In what percentage of the inspections is there a non-null code?

```python
violation_pct = ...
```

**3.** Tabulate (i.e., create a series) by city the percentage of inspections that have a violation. 

```python
violation_bycity = ...
```

```python
grader.check("inspect-bycity")
```

## Births



**1.** Load the dataset `Births.csv` as a data frame. (This dataset describes all live births in Delaware over several years.)

```python
births = ...
```

```python
grader.check("births-load")
```

<!-- BEGIN QUESTION -->

**2.** Make a stacked histogram plot showing the distribution of birth weights, separated into columns according to the mother's marital status, and using color to show the gender.

```python
...

```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

**3.** Are there strong correlations between the birth weight and gestation time? To explore this question, extend the frame with dummy variables for `BIRTH WEIGHT` and `GESTATION`, and then compute pairwise correlation coefficients. Briefly interpret the results, and note whether the findings are as expected or surprising.

```python

```

<!-- END QUESTION -->

## Green energy rebates


Here we are loading the dataset `Green_Energy_Program_Grants.csv` as a frame called `grants`. 

```python
grants = pd.read_csv("Green_Energy_Program_Grants.csv")
```


There is some sloppiness in the *Sector* column, because some values are capitalized while others are not. 


```python
grants["Sector"].unique()
```

**1.** Replace the original frame by one where the lowercase forms are replaced by the capitalized forms.

```python
grants = ...
```

```python
grader.check("grants-load")
```

**2.** Add a column called *percent* to the frame for the percentage of the total cost that was rebated. Create a series that includes its mean, standard deviation, and quartile values.

```python
...
stats = ...
```

```python
grader.check("grants-capital")
```

**3.** Find the Spearman correlation coefficient between total system costs and total rebate. The result should just be a single number (e.g., no self-correlations).

```python
corr = ...

```

```python
grader.check("grants-corr")
```

<!-- BEGIN QUESTION -->

**4.** Make a scatter plot of total rebate as a function of total cost, separating into columns by *Sector*.

```python
...
```

<!-- END QUESTION -->

**5.** Find the Spearman correlation coefficient between total system costs and total rebate within each type of *Sector*. (The result should be a series of 5 positive values.)

```python
corr_bysector = ...

```

```python
pd.Series(C.values,index=[i[0] for i in C.index])
```

```python
grader.check("grants-corr_sector")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Select *Kernel/Restart & Run All*, then save, then run this export cell again. Submit by pushing the resulting zip file to your GitHub assignment repo.

```python
grader.export(pdf=False, force_save=True)
```

Tests for birds-Dec (points: None)


Tests for birds-canada (points: None)
```
>>> assert len(canada) == 11
```
failure_message: The length of the value is incorrect.
```
>>> assert canada[3.0] == 1968090
```
failure_message: Zone 3 value is incorrect.


Tests for birds-load (points: None)
```
>>> assert birds.shape == (2045,37)
```
failure_message: The shape of the value is incorrect.


Tests for birds-pearson (points: None)
```
>>> 
>>> assert duckduck.shape==(4,4)
```
failure_message: The shape is incorrect.
```
>>> assert np.isclose(duckduck.loc['Mallard','Wood Duck'],-0.0058356786949663115)
```
failure_message: At least one correlation value is incorrect.


Tests for birds-spearman (points: None)
```
>>> 
>>> assert duckspear.shape==(4,4)
```
failure_message: The shape is incorrect.
```
>>> assert np.isclose(duckspear.loc['Mallard','Wood Duck'],0.16751354618493336)
```
failure_message: At least one correlation value is incorrect.


Tests for birds-zone3 (points: None)
```
>>> for bird in ['Atlantic Brant', 'Canada Goose', 'Snow Goose',
...        'Greater White-fronted Goose', 'American Black Duck', 'Mallard',
...        'Northern Pintail', 'American Green-winged Teal', 'Blue-winged Teal',
...        'Gadwall', 'American Wigeon', 'Northern Shoveler', 'Wood Duck',
...        'Canvasback', 'Redhead', 'Scaup', 'Ring-necked Duck',
...        'Common Merganser', 'Red-breasted Merganser', 'Hooded Merganser',
...        'Mergansers', 'Ruddy Duck', 'White-winged Scoter', 'Surf Scoter',
...        'Black Scoter', 'Scoters', 'Bufflehead', 'Common Goldeneye',
...        'Long-tailed Duck', 'Mute Swan', 'Tundra Swan', 'American Coot']:
...        assert bird in zone3.index
```
failure_message: Missing at least one species.


Tests for births-load (points: None)
```
>>> assert births.shape==(96457,18)
```


Tests for grants-capital (points: None)
```
>>> for col in ["mean","std","25%","50%","75%"]:
...     assert col in stats.index
```
failure_message: Use names "mean", "std", "25%", "50%", "75%" in the index.
```
>>> assert np.isclose(grants.percent.mean(),18.192261185675047)
```
failure_message: Values in the percent column may be incorrect.


Tests for grants-corr (points: None)
```
>>> assert type(corr) in [float,np.float64]
```
failure_message: The result should be a single floating-point number.


Tests for grants-corr_sector (points: None)
```
>>> assert type(corr_bysector) == pd.Series
```
failure_message: Result must be of type Series.
```
>>> assert np.isclose(corr_bysector["Non_Residential"],0.8185436634413635)
```
failure_message: At least one value is missing or wrong.


Tests for grants-load (points: None)
```
>>> assert grants.shape==(5187,15)
```
failure_message: The data frame has the wrong shape.
```
>>> assert list(grants.Sector.unique())==['Residential', 'Non_Residential', 'Solar Lease', 'Solar PPA',
...        'Non_Profit']
```
failure_message: The *Sector* column does not have all correct possible values.


Tests for inspect-bycity (points: None)
```
>>> assert len(violation_bycity)==43
```
```
>>> assert np.isclose(violation_bycity["WILMINGTON"],2.304147465437788)
```


Tests for inspect-load (points: None)
```
>>> assert inspects.shape == (1647,9)
```
```
>>> assert inspects.loc[1000,"Food Establishment Zip Code"] == 19971
```


Tests for inspect-nan (points: None)
